# CMT318: Lab 1 - Python Fundamentals with Basic Text Processing (Autumn 2024)

This notebook will serve to practice important Python concepts together with some basic preprocessing techniques. It has been extracted from [UW's CSE 447 / CSE M 547 Project 0](https://drive.google.com/file/d/1hiZ278EJCRp0iJntYO4c2tyZOzI4SLSP/view) originally written by [Kabir Ahuja](https://kabirahuja2431.github.io/).

This noteboook covers basic python, such as:

* Basic Data Types
* Functions
* Containers or Data structures: Lists, Dictionaries, Sets, Tuples,
* Classes

Three exercises related to text preprocessing are also included. Each exercise grows in complexity. Students who are already familiar with Python can go straight to the exercises if they prefer.

**IMPORTANT: Save a copy of the notebook before working on it. You can also download it if you prefer**

### Basic data types

#### Numbers

Integers and floats work as you would expect from other languages:

```python
x = 3
print(x, type(x))
```

```python
print(x + 1)  # Addition
print(x - 1)  # Subtraction
print(x * 2)  # Multiplication
print(x**2)  # Exponentiation
```

```python
x += 1
print(x)
x *= 2
print(x)
```

Note that unlike many languages, Python does not have unary increment (x++) or decrement (x--) operators.

Python also has built-in types for long integers and complex numbers; you can find all of the details in the [documentation](https://docs.python.org/3.7/library/stdtypes.html#numeric-types-int-float-long-complex).

#### Booleans

Python implements all of the usual operators for Boolean logic, but uses English words rather than symbols (`&&`, `||`, etc.):

```python
t, f = True, False
print(type(t))
```

Now we let's look at the operations:

```python
print(t and f)  # Logical AND;
print(t or f)  # Logical OR;
print(not t)  # Logical NOT;
print(t != f)  # Logical XOR;
```

#### Strings

```python
hello = "hello"  # String literals can use single quotes
world = "world"  # or double quotes; it does not matter
print(hello, len(hello))
```

```python
hw = hello + " " + world  # String concatenation
print(hw)
```

```python
hw12 = f"{hello} {world} {12}"  # string formatting
print(hw12)
```

String objects have a bunch of useful methods; for example:

```python
s = "hello"
print(s.capitalize())  # Capitalize a string
print(s.upper())  # Convert a string to uppercase; prints "HELLO"
print(s.rjust(7))  # Right-justify a string, padding with spaces
print(s.center(7))  # Center a string, padding with spaces
print(s.replace("l", "(ell)"))  # Replace all instances of one substring with another
print("  world ".strip())  # Strip leading and trailing whitespace
```

You can find a list of all string methods in the [documentation](https://docs.python.org/3.7/library/stdtypes.html#string-methods).

### Functions

Python functions are defined using the `def` keyword. For example:

```python
def sign(x):
    if x > 0:
        return "positive"
    elif x < 0:
        return "negative"
    else:
        return "zero"


print(sign(3))
print(sign(-3))
print(sign(0))
```

Output:
```
positive
negative
zero
```

We will often define functions to take optional keyword arguments, like this:

```python
def hello(name, loud=False):
    if loud:
        print(f"HELLO, {name.upper()}")
    else:
        print(f"Hello, {name}!")


hello("Bob")
hello("Fred", loud=True)
```

Output:
```
Hello, Bob!
HELLO, FRED
```

### Exercise 1: Basic Text Processing Using String Methods (Level 1)

You will now use what we just learned to implement some basic text processing methods in Python. Text processing comes very much in handy when training machine learning models on natural language data. Since text processing is typically done before training of the models, it is also commonly referred to as *preprocessing*. For this exercise you will perform the following basic text processing steps:

- convert the text to lower case
- remove extra spaces (left trailing, right trailing, or in between the words)
- remove punctuations

Note that, while more advanced libraries like regular expressions (re) or natural language toolkit (NLTK) can also be used for implementing these methods, it is also straightforward to use string methods to carry out these basic operations. Try implementing these functions using string methods first, and then with functions from NLTK.

**Important Note: Remove `raise NotImplementedError()` when you write your code in the functions**

<!-- grade_id: cell-243665322a6645a5, points: 0, grade: false, solution: true -->
```python
def to_lowercase(text: str) -> str:
    """
    Convert a string to lowercase

    E.g. "Hello" -> "hello"

    Input:
        - text: string
    Output:
        - string

    """
    lowercase_text = None
    # YOUR CODE HERE
    raise NotImplementedError()
    return lowercase_text


def remove_extra_spaces(text: str) -> str:
    """
    Remove extra spaces from a string.

    E.g.: "  This is a    string   " -> "This is a string"

    Input:
        - text: string
    Output:
        - string

    """
    text_without_extra_spaces = None

    # YOUR CODE HERE
    raise NotImplementedError()

    return text_without_extra_spaces

def remove_punctuations(text):
    """
    Remove punctuations from a string

    E.g.: "Hello! How are you?" -> "Hello How are you"

    Input:
        - text: string
    Output:
        - string

    """
    # we can get a list of punctuations from the string module
    import string
    punctuations = string.punctuation

    text_without_punctuations = None

    # YOUR CODE HERE
    raise NotImplementedError()

    return text_without_punctuations


def text_processing(text):
    """
    Process a text by converting it to lowercase, removing extra spaces and removing punctuations

    Input:
        - text: string
    Output:
        - string

    """
    processed_text = None

    # YOUR CODE HERE
    raise NotImplementedError()

    return processed_text
```

<!-- grade_id: cell-0d71367e7eadb2e5, points: 1, grade: true, solution: false -->
```python
def test_text_processing():
    assert text_processing("  This is a    string   ") == "this is a string"
    assert text_processing("Hello! How are you?") == "hello how are you"
    assert text_processing("  This is a    string   with punctuations! ") == "this is a string with punctuations"
    print("All tests pass")

test_text_processing()
```

### Containers

Python includes several built-in container types: lists, dictionaries, sets, and tuples.

#### Lists

A list is the Python equivalent of an array, but is resizeable and can contain elements of different types:

```python
xs = [3, 1, 2]  # Create a list
print(xs, xs[2])
print(xs[-1])  # Negative indices count from the end of the list; prints "2"
```

```python
xs[2] = "foo"  # Lists can contain elements of different types
print(xs)
```

```python
xs.append("bar")  # Add a new element to the end of the list
print(xs)
```

```python
x = xs.pop()  # Remove and return the last element of the list
print(x, xs)
```

As usual, you can find all the gory details about lists in the [documentation](https://docs.python.org/3.7/tutorial/datastructures.html#more-on-lists).

#### Slicing

In addition to accessing list elements one at a time, Python provides concise syntax to access sublists; this is known as slicing:

```python
nums = list(range(5))  # range is a built-in function that creates a list of integers
print(nums)  # Prints "[0, 1, 2, 3, 4]"
print(nums[2:4])  # Get a slice from index 2 to 4 (exclusive); prints "[2, 3]"
print(nums[2:])  # Get a slice from index 2 to the end; prints "[2, 3, 4]"
print(nums[:2])  # Get a slice from the start to index 2 (exclusive); prints "[0, 1]"
print(nums[:])  # Get a slice of the whole list; prints ["0, 1, 2, 3, 4]"
print(nums[:-1])  # Slice indices can be negative; prints ["0, 1, 2, 3]"
nums[2:4] = [8, 9]  # Assign a new sublist to a slice
print(nums)  # Prints "[0, 1, 8, 9, 4]"
```

#### Loops

You can loop over the elements of a list like this:

```python
animals = ["cat", "dog", "monkey"]
for animal in animals:
    print(animal)
```

If you want access to the index of each element within the body of a loop, use the built-in `enumerate` function:

```python
animals = ["cat", "dog", "monkey"]
for idx, animal in enumerate(animals):
    print(f"#{idx+1}: {animal}")
```

#### List comprehensions:
When programming, frequently we want to transform one type of data into another. As a simple example, consider the following code that computes square numbers:

```python
nums = [0, 1, 2, 3, 4]
squares = []
for x in nums:
    squares.append(x**2)
print(squares)
```

You can make this code simpler using a list comprehension:

```python
nums = [0, 1, 2, 3, 4]
squares = [x**2 for x in nums]
print(squares)
```

List comprehensions can also contain conditions:

```python
nums = [0, 1, 2, 3, 4]
even_squares = [x**2 for x in nums if x % 2 == 0]
print(even_squares)
```

#### Dictionaries

A dictionary stores (key, value) pairs, similar to a `Map` in Java or an object in Javascript. You can use it like this:

```python
d = {"cat": "cute", "dog": "furry"}  # Create a new dictionary with some data
print(d["cat"])  # Get an entry from a dictionary; prints "cute"
print("cat" in d)  # Check if a dictionary has a given key; prints "True"
```

```python
d["fish"] = "wet"  # Set an entry in a dictionary
print(d["fish"])  # Prints "wet"
```

```python
print(d["monkey"])  # KeyError: 'monkey' not a key of d
```

```python
print(d.get("monkey", "N/A"))  # Get an element with a default; prints "N/A"
print(d.get("fish", "N/A"))  # Get an element with a default; prints "wet"
```

```python
del d["fish"]  # Remove an element from a dictionary
print(d.get("fish", "N/A"))  # "fish" is no longer a key; prints "N/A"
```

You can find all you need to know about dictionaries in the [documentation](https://docs.python.org/3/library/stdtypes.html#mapping-types-dict).

It is easy to iterate over the keys in a dictionary:

```python
d = {"person": 2, "cat": 4, "spider": 8}
for animal, legs in d.items():
    print(f"A {animal} has {legs} legs")
```

Dictionary comprehensions: These are similar to list comprehensions, but allow you to easily construct dictionaries. For example:

```python
nums = [0, 1, 2, 3, 4]
even_num_to_square = {x: x**2 for x in nums if x % 2 == 0}
print(even_num_to_square)
```

#### Sets
A set is an unordered collection of distinct elements. As a simple example, consider the following:

```python
animals = {"cat", "dog"}
print("cat" in animals)  # Check if an element is in a set; prints "True"
print("fish" in animals)  # prints "False"
```

```python
animals.add("fish")  # Add an element to a set
print("fish" in animals)
print(len(animals))  # Number of elements in a set;
```

```python
animals.add("cat")  # Adding an element that is already in the set does nothing
print(len(animals))
animals.remove("cat")  # Remove an element from a set
print(len(animals))
```

_Loops_: Iterating over a set has the same syntax as iterating over a list; however since sets are unordered, you cannot make assumptions about the order in which you visit the elements of the set:

```python
animals = {"cat", "dog", "fish"}
for idx, animal in enumerate(animals):
    print(f"#{idx + 1}: {animal}")
```

Set comprehensions: Like lists and dictionaries, we can easily construct sets using set comprehensions:

```python
from math import sqrt

print({int(sqrt(x)) for x in range(30)})
```

#### Tuples

A tuple is an (immutable) ordered list of values. A tuple is in many ways similar to a list; one of the most important differences is that tuples can be used as keys in dictionaries and as elements of sets, while lists cannot. Here is a trivial example:

```python
d = {(x, x + 1): x for x in range(10)}  # Create a dictionary with tuple keys
t = (5, 6)  # Create a tuple
print(type(t))
print(d[t])
print(d[(1, 2)])
```

```python
t[0] = 1
```

### Exercise 2: Word Tokenization and Converting Tokens to IDs (Level 2)

When working with text data most Machine Learning / Deep Learning algorithms treat text as a sequence of smaller units, which are also called tokens. While there exists different levels of granularities on what constitutes as a token, for this notebook we will focus on the classical approach of word tokenization i.e. the text is split into a sequence of words.

E.g. ``` "This is a sentence about to be word tokenized" -> ["This", "is", "a", "sentence", "about", "to", "be", "word", "tokenized"]```

Tokenization is usually followed by a step converting the tokens into a set of input ids. As you will learn during the course, machines do not have semantic knowledge of words built in, and learning algorithms treat these words as numbers identifying (IDs) each word or token. These IDs are typically assigned according to the index of a word in the vocabulary, which is a list of all unique words in the training corpus.

E.g. For a vocabulary: ```vocab = ["A", a", "cat", "mat", "on", "sits"]```, the sentence ```"a cat sits on a mat"``` with tokens ```["A", "cat", "sits", "on", "a", "mat"]```, will get converted to the list of ids: ```[0, 2, 5, 4, 1, 3]```

Below you will implement the necessary functionality to implement word tokenization and conversion of words into ids. Specifically you will implement three functions:

- `word_tokenize` : Converts a text (represented as python string) into a sequence of words
- `fit_vocab`: Constructs vocabulary i.e. list of unique tokens from a corpus (a list of text documents)
- `convert_token_to_ids`: Converts a list of tokens to ids

You will be using functionalities from different containers like lists, sets, and dictionaries, which you just learned. Like previous exercise, we recommend you to avoid using any external libraries to do this exercise to maximize your learning of the underlying concepts.

**Important Note: Remove `raise NotImplementedError()` when you write your code in the functions**


<!-- grade_id: cell-14520531a5c2169c, points: 0, grade: false, solution: true -->
```python
def word_tokenize(text: str) -> list:
    """
    Tokenize a text into words. You can assume that no punctuations are present in the text.

    E.g.: "we all live in a yellow submarine" -> ["we", "all", "live", "in", "a", "yellow", "submarine"]

    Input:
        - text: string
    Output:
        - list of strings

    """
    words = None
    # YOUR CODE HERE
    raise NotImplementedError()
    return words
```

<!-- grade_id: cell-da4dc5898efa830f, points: 0.5, grade: true, solution: false -->
```python
def test_word_tokenize():
    assert word_tokenize("we all live in a yellow submarine") == [
        "we",
        "all",
        "live",
        "in",
        "a",
        "yellow",
        "submarine",
    ]
    assert word_tokenize("This is a sentence about to be word tokenized?") == [
        "This",
        "is",
        "a",
        "sentence",
        "about",
        "to",
        "be",
        "word",
        "tokenized?",
    ]
    print("All tests pass")


test_word_tokenize()
```

<!-- grade_id: cell-e2f0d372919b87c4, points: 0, grade: false, solution: true -->
```python
def fit_vocab(corpus):
    """
    Creates a vocabulary from a corpus i.e. a list of documents.


    Input:
        - corpus: list of strings
    Output:
        - List of unique words in the corpus (sorted in alphabetical order)
        - Dictionary mapping each word to its index in the vocabulary

    Important: Remember to sort the vocabulary before using it to construct the dictionary mapping
    """

    vocab = None
    word_to_idx = None

    # YOUR CODE HERE
    raise NotImplementedError()

    return vocab, word_to_idx
```

<!-- grade_id: cell-da20f6bfa22f808b, points: 0.5, grade: true, solution: false -->
```python
def test_fit_vocab():
    corpus = [
        "We all live in a yellow submarine",
        "I am the walrus",
        "Yellow Submarine",
    ]
    vocab, word_to_idx = fit_vocab(corpus)
    assert vocab == [
        "I",
        "Submarine",
        "We",
        "Yellow",
        "a",
        "all",
        "am",
        "in",
        "live",
        "submarine",
        "the",
        "walrus",
        "yellow",
    ]
    assert word_to_idx == {
        "I": 0,
        "Submarine": 1,
        "We": 2,
        "Yellow": 3,
        "a": 4,
        "all": 5,
        "am": 6,
        "in": 7,
        "live": 8,
        "submarine": 9,
        "the": 10,
        "walrus": 11,
        "yellow": 12,
    }

    print("All tests pass")


test_fit_vocab()
```

<!-- grade_id: cell-437d01662d031940, points: 0, grade: false, solution: true -->
```python
def convert_tokens_to_ids(tokens, word_to_idx):
    """
    Convert a list of tokens to a list of token ids using the word_to_idx dictionary.

    Input:
        - tokens: list of strings
        - word_to_idx: dictionary mapping words to their ids
    Output:
        - list of integers
    """

    token_ids = None
    # YOUR CODE HERE
    raise NotImplementedError()
    return token_ids
```

<!-- grade_id: cell-1c0ee37757d295c9, points: 0.5, grade: true, solution: false -->
```python
def test_convert_tokens_to_ids():
    word_to_idx = {
        "I": 0,
        "Submarine": 1,
        "We": 2,
        "Yellow": 3,
        "a": 4,
        "all": 5,
        "am": 6,
        "in": 7,
        "live": 8,
        "submarine": 9,
        "the": 10,
        "walrus": 11,
        "yellow": 12,
    }
    tokens = ["We", "all", "live", "in", "a", "yellow", "submarine"]
    assert convert_tokens_to_ids(tokens, word_to_idx) == [2, 5, 8, 7, 4, 12, 9]
    print("All tests pass")


test_convert_tokens_to_ids()
```

### Classes

The syntax for defining classes in Python is straightforward:

```python
class Greeter:

    # Constructor
    def __init__(self, name):
        self.name = name  # Create an instance variable

    # Instance method
    def greet(self, loud=False):
        if loud:
            print("HELLO, {}".format(self.name.upper()))
        else:
            print("Hello, {}!".format(self.name))


g = Greeter("Fred")  # Construct an instance of the Greeter class
g.greet()  # Call an instance method; prints "Hello, Fred"
g.greet(loud=True)  # Call an instance method; prints "HELLO, FRED!"
```

### Exercise 3: Implementing the Full Tokenization Pipeline from Scratch (Level 3)

We will now take everything we learned to implement a typical tokenization pipeline very similar to the one used in [Hugging Face Transformers library](https://huggingface.co/docs/tokenizers/en/pipeline). In particular, you will implement the class `WordTokenizer` with the following methods:

- `normalizer`: Applies the text processing techniques to the input text (we will use the same 3 techniques we implemented before; we use a different terminology i.e. normalizer instead of processing to be close to the huggingface naming convention)
- `tokenize`: Splits the text into a list of words
- `convert_tokens_to_ids`: Converts a list of tokens to a list of ids
- `convert_ids_to_tokens`: Converts a list of ids to a list of tokens
- `encode`: Applies `tokenize` and `convert_tokens_to_ids` methods in succession to an input text
- `decode`: Converts the list of ids back to the text form
- `train`: Trains the tokenizer, which for the word tokenizer means simply fitting the vocabulary over the corpus
  
**Important Note: Remove `raise NotImplementedError()` when you write your code in the functions**

<!-- grade_id: cell-66b4de2616f3e708, points: 0, grade: false, solution: true -->
```python
class WordTokenizer:

    def __init__(self):
        """
        Constructor for the WordTokenizer class.

        Initialize `vocab` attribute as an empty list and `word_to_idx` attribute as an empty dictionary.
        """

        # YOUR CODE HERE
        raise NotImplementedError()

    def normalizer(self, text):
        """
        Normalizes the input text by converting it to lowercase, removing extra spaces, and removing punctuations.

        Input:
            - text: string
        Output:
            - string

        """

        normalized_text = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return normalized_text

    def tokenize(self, text):
        """
        Tokenizes the input text into words.

        Input:
            - text: string
        Output:
            - list of strings

        """

        tokens = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return tokens

    def convert_tokens_to_ids(self, tokens):
        """
        Convert a list of tokens to a list of token ids using the word_to_idx dictionary.

        Input:
            - tokens: list of strings
        Output:
            - list of integers

        """

        # Throw an error if the word_to_idx dictionary is empty
        if not self.word_to_idx:
            raise ValueError("Tokenizer has not been fit on a vocabulary yet. Call `train` method with a corpus first!")

        token_ids = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return token_ids

    def convert_ids_to_tokens(self, token_ids):
        """
        Convert a list of token ids to a list of tokens using the idx_to_word dictionary.

        Input:
            - token_ids: list of integers
        Output:
            - list of strings

        """

        # Throw an error if the vocabulary is empty
        if self.vocab == []:
            raise ValueError("Tokenizer has not been fit on a vocabulary yet. Call `train` method with a corpus first!")

        tokens = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return tokens

    def encode(self, text):
        """
        Encodes the input text into token ids. Do not forget to normalize the text before tokenizing it.

        Input:
            - text: string
        Output:
            - list of integers

        """

        token_ids = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return token_ids

    def decode(self, token_ids):
        """
        Decodes the input token ids into text.

        Input:
            - token_ids: list of integers
        Output:
            - string

        """
        text = None
        # YOUR CODE HERE
        raise NotImplementedError()
        return text

    def train(self, corpus):
        """
        Trains the tokenizer on a corpus i.e. a list of documents, filling in values of `self.vocab` and `self.word_to_idx`.

        Input:
            - corpus: list of strings

        Note: Do not forget to normalize the text before tokenizing it.
        """

        # YOUR CODE HERE
        raise NotImplementedError()

```

<!-- grade_id: cell-65defcf000eeba30, points: 1.5, grade: true, solution: false -->
```python
def test_WordTokenizer():
    tokenizer = WordTokenizer()
    corpus = [
        "We all live in a yellow submarine",
        "I am the walrus",
        "Yellow Submarine",
    ]
    tokenizer.train(corpus)
    token_ids = tokenizer.encode("We all live in a yellow submarine")

    assert token_ids == [9, 1, 5, 4, 0, 10, 6]

    text = tokenizer.decode(token_ids)
    assert text == "we all live in a yellow submarine"

    print("All tests pass")


test_WordTokenizer()
```

```python

```